Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name below:

```OCaml
let name = "Anumala Venu Madhava Reddy"
let rollno = "CS18B051"
```

Output:
```
val name : string = "Anumala Venu Madhava Reddy"

val rollno : string = "CS18B051"
```

## Important notes about grading:

1. **Compiler errors:** All code you submit must compile. Programs that do not compile will probably receive an automatic zero. If you are having trouble getting your assignment to compile, please visit consulting hours. If you run out of time, it is better to comment out the parts that do not compile, than hand in a more complete file that does not compile.
2. **Late assignments:** Please carefully review the course website's policy on late assignments, as all assignments handed in after the deadline will be considered late. Verify on moodle that you have submitted the correct version, before the deadline. Submitting the incorrect version before the deadline and realizing that you have done so after the deadline will be counted as a late submission.

# Lambda Calculus Interpreter

In this assignment, you will implement lambda calculus interpreters that use different reduction strategies. The abstract syntax tree (AST) for lambda expressions is the one that we have seen in class:


```ocaml
type expr = 
  | Var of string
  | Lam of string * expr
  | App of expr * expr
```


You are provided a parser function `parse_string` that converts a string to this AST and a printer function `string_of_expr` that converts the AST to string. For example, 

```OCaml
#use "init.ml"
open Syntax
let parse_string = Lambda_parse.parse_string
let string_of_expr = string_of_expr

let _ = parse_string "(\\x.x) (\\y.y)"
let _ = string_of_expr (App (Var "x",Lam("y",App(Var "y", Var "x"))))
```

Output:
```
val alpha_equiv : Syntax.expr -> Syntax.expr -> bool = <fun>
val parse_string : string -> Syntax.expr = <fun>

val string_of_expr : Syntax.expr -> string = <fun>

- : Syntax.expr = App (Lam ("x", Var "x"), Lam ("y", Var "y"))

- : string = "x (λy.y x)"
```

You will need some helper functions to operate over sets. Since we have not studied set data structure in OCaml. We will use lists instead and implement set functionality on top of lists. You can use the functions from the OCaml list standard library for this assignment. The documentation for OCaml list module is available in the [OCaml manual](https://caml.inria.fr/pub/docs/manual-ocaml/libref/List.html). 

## Problem 1

Implement a function

```ocaml
mem : 'a -> 'a list -> bool
```

`mem e l` returns `true` if the element `e` is present in the list. Otherwise, it returns false. 

<!-- grade_id: cell-206225a8fc096ef3, points: 0, grade: false, solution: true -->
```OCaml
let rec mem e l acc = 
    match l with 
    |[] -> acc
    |h::t -> if e=h then true else mem e t acc
   
let mem e l = mem e l false 
```

Output:
```
val mem : 'a -> 'a list -> bool -> bool = <fun>

val mem : 'a -> 'a list -> bool = <fun>
```

<!-- grade_id: cell-625148cc24b57f40, points: 5, grade: true, solution: false -->
```OCaml
assert (mem "b" ["a";"b";"c"] = true);
assert (mem "x" ["a";"b";"c"] = false)
```

Output:
```
- : unit = ()
```

## Problem 2

Implement a function

```ocaml
remove : 'a -> 'a list -> 'a list
```

`remove e l` returns a list `l'` with all the element in `l` except `e`. `remove` also preserves the order of the elements not removed. If `e` is not present in `l`, then return `l`.

<!-- grade_id: cell-3d615242df1b1ad1, points: 0, grade: false, solution: true -->
```OCaml
let rec rev_list l acc = 
  match l with
  | [] -> acc
  | hd::tl -> rev_list tl (hd::acc)

let rev_list l = rev_list l []

let rec remove e l l1 =
    match l with
    |[] -> rev_list l1
    |h::t -> if e=h then remove e t l1 else (remove e t (h::l1))
    
let remove e l = remove e l []
```

Output:
```
val rev_list : 'a list -> 'a list -> 'a list = <fun>

val rev_list : 'a list -> 'a list = <fun>

val remove : 'a -> 'a list -> 'a list -> 'a list = <fun>

val remove : 'a -> 'a list -> 'a list = <fun>
```

<!-- grade_id: cell-815b5c339ce6eec3, points: 5, grade: true, solution: false -->
```OCaml
assert (remove "b" ["a";"b";"c"] = ["a";"c"]);
assert (remove "x" ["a";"b";"c"] = ["a";"b";"c"])
```

Output:
```
- : unit = ()
```

## Problem 3

Implement a function

```ocaml
union : string list -> string list -> string list
```

`union l1 l2` performs set union of elements in `l1` and `l2`. The elements in the result list `l` must be lexicographically sorted. Hint: You may want to use the functions `List.sort` and `remove_stutter` from assignment 1 to implement union. Here is an example of using `List.sort`. 

```OCaml
assert (List.sort String.compare ["x";"a";"b";"m"] = ["a";"b";"m";"x"])
```

Output:
```
- : unit = ()
```

<!-- grade_id: cell-e498f15f8517e45a, points: 0, grade: false, solution: true -->
```OCaml
let rec rev_list l acc = 
  match l with
  | [] -> acc
  | hd::tl -> rev_list tl (hd::acc)

let rev_list l = rev_list l []

let rec remove_fun l1 x acc = 
    match l1 with
    |[] -> rev_list acc
    |hd1::tl1 -> if hd1=x then remove_fun tl1 hd1 acc
                 else remove_fun tl1 hd1 (hd1::acc)


let remove_stutter l =
    match l with
    |[]->l
    |hd::t1-> remove_fun t1 hd (hd::[])
    
let rec union l1 l2 = 
    match l2 with 
    |[]->remove_stutter (List.sort String.compare l1)
    |h::t-> union (l1@[h]) t
```

Output:
```
val rev_list : 'a list -> 'a list -> 'a list = <fun>

val rev_list : 'a list -> 'a list = <fun>

val remove_fun : 'a list -> 'a -> 'a list -> 'a list = <fun>

val remove_stutter : 'a list -> 'a list = <fun>

val union : String.t list -> String.t list -> String.t list = <fun>
```

<!-- grade_id: cell-42c81525ae442418, points: 5, grade: true, solution: false -->
```OCaml
assert (union ["a"; "c"; "b"] ["d"; "b"; "x"; "a"] = ["a"; "b"; "c"; "d"; "x"])
```

Output:
```
- : unit = ()
```

## Problem 4

Implement a function

```ocaml
add : 'a -> 'a list -> 'a list
```

`add e l` does a set addition of element `e` to list `l` and returns a list. The resultant list is sorted. 

<!-- grade_id: cell-de42368283c8797d, points: 0, grade: false, solution: true -->
```OCaml
let rec add e l l1= 
    match l with
    |[] -> List.sort compare (l1@[e])
    |h::t -> if e=h then add e t l1 else add e t (l1@[h])
    
let add e l = add e l []
```

Output:
```
val add : 'a -> 'a list -> 'a list -> 'a list = <fun>

val add : 'a -> 'a list -> 'a list = <fun>
```

<!-- grade_id: cell-635dea866c593bca, points: 5, grade: true, solution: false -->
```OCaml
assert (add "b" ["a";"c"] = ["a";"b";"c"]);
assert (add "a" ["c"; "a"] = ["a";"c"])
```

Output:
```
- : unit = ()
```

## Substitution

At the heart of reducing lambda expressions is substitution. Recall from the lecture that substitution requires us to generate fresh variable names that is different from every other name used in the current context. We will use the following helper function to generate fresh names.

```OCaml
let r = ref 0                                                                       
                                                                                    
let fresh s =                                                                       
  let v = !r in                                                                     
  r := !r + 1;                                                                      
  s ^ (string_of_int v)
```

Output:
```
val r : int ref = {contents = 0}

val fresh : string -> string = <fun>
```

It uses mutability features of OCaml which we will study in later lectures. You can use the `fresh` function as follows:

```OCaml
let a = fresh "a"
let b = fresh "b"
```

Output:
```
val a : string = "a0"

val b : string = "b1"
```

## Problem 5

Implement a function 

```ocaml
free_variables: expr -> string list
```

that returns the free variables in the given lambda term.

<!-- grade_id: cell-8ad0bb22f3fca6ed, points: 0, grade: false, solution: true -->
```OCaml
let rec rev_list l acc = 
  match l with
  | [] -> acc
  | hd::tl -> rev_list tl (hd::acc)

let rev_list l = rev_list l []

let rec remove e l l1 =
    match l with
    |[] -> rev_list l1
    |h::t -> if e=h then remove e t l1 else (remove e t (h::l1))
    
let remove e l = remove e l []

let rec remove_fun l1 x acc = 
    match l1 with
    |[] -> rev_list acc
    |hd1::tl1 -> if hd1=x then remove_fun tl1 hd1 acc
                 else remove_fun tl1 hd1 (hd1::acc)


let remove_stutter l =
    match l with
    |[]->l
    |hd::t1-> remove_fun t1 hd (hd::[])
    
let rec union l1 l2 = 
    match l2 with 
    |[]->remove_stutter (List.sort String.compare l1)
    |h::t-> union (l1@[h]) t
    
let rec free_variables e l =
    match e with
    |Lam(x,t1) -> remove x (free_variables t1 l)
    |App(t1, t2) -> union (free_variables t1 l) (free_variables t2 l) 
    |Var x -> l@[x]
    
let free_variables e = free_variables e []

```

Output:
```
val rev_list : 'a list -> 'a list -> 'a list = <fun>

val rev_list : 'a list -> 'a list = <fun>

val remove : 'a -> 'a list -> 'a list -> 'a list = <fun>

val remove : 'a -> 'a list -> 'a list = <fun>

val remove_fun : 'a list -> 'a -> 'a list -> 'a list = <fun>

val remove_stutter : 'a list -> 'a list = <fun>

val union : String.t list -> String.t list -> String.t list = <fun>

val free_variables : Syntax.expr -> String.t list -> String.t list = <fun>

val free_variables : Syntax.expr -> String.t list = <fun>
```

<!-- grade_id: cell-67d38d859ec6173a, points: 20, grade: true, solution: false -->
```OCaml
assert (free_variables (parse_string "\\x.x") = []);
assert (free_variables (parse_string "\\x.y") = ["y"])
```

Output:
```
- : unit = ()
```

## Problem 6

Implement the function

```ocaml
substitute : expr -> string -> expr -> expr
```

where `substitute e x v` does `e[v/x]`. For renaming `x` in `Lam(x,y)` with a fresh name, use `Lam (fresh x, ...)`.

<!-- grade_id: cell-c659e5c57a377290, points: 0, grade: false, solution: true -->
```OCaml
let rec mem e l acc = 
    match l with 
    |[] -> acc
    |h::t -> if e=h then true else mem e t acc
   
let mem e l = mem e l false 

let rec rev_list l acc = 
  match l with
  | [] -> acc
  | hd::tl -> rev_list tl (hd::acc)

let rev_list l = rev_list l []

let rec remove e l l1 =
    match l with
    |[] -> rev_list l1
    |h::t -> if e=h then remove e t l1 else (remove e t (h::l1))
    
let remove e l = remove e l []

let rec remove_fun l1 x acc = 
    match l1 with
    |[] -> rev_list acc
    |hd1::tl1 -> if hd1=x then remove_fun tl1 hd1 acc
                 else remove_fun tl1 hd1 (hd1::acc)


let remove_stutter l =
    match l with
    |[]->l
    |hd::t1-> remove_fun t1 hd (hd::[])
    
let rec union l1 l2 = 
    match l2 with 
    |[]->remove_stutter (List.sort String.compare l1)
    |h::t-> union (l1@[h]) t
    
let rec free_variables e l =
    match e with
    |Lam(x,t1) -> remove x (free_variables t1 l)
    |App(t1, t2) -> union (free_variables t1 l) (free_variables t2 l) 
    |Var x -> l@[x]
    
let free_variables e = free_variables e []

let rec substitute expr a b = 
    match expr with 
    |Var x -> if a = x then  b else Var x
    |Lam(x,t1) -> if a=x then Lam(x,t1) else 
         if mem x (free_variables b) = false then Lam(x,substitute t1 a b) 
         else  let w = fresh (x) in substitute (  Lam(w,substitute t1 x (Var w))) a b 
    |App(t1,t2) -> App(substitute t1 a b,substitute t2 a b)
```

Output:
```
val mem : 'a -> 'a list -> bool -> bool = <fun>

val mem : 'a -> 'a list -> bool = <fun>

val rev_list : 'a list -> 'a list -> 'a list = <fun>

val rev_list : 'a list -> 'a list = <fun>

val remove : 'a -> 'a list -> 'a list -> 'a list = <fun>

val remove : 'a -> 'a list -> 'a list = <fun>

val remove_fun : 'a list -> 'a -> 'a list -> 'a list = <fun>

val remove_stutter : 'a list -> 'a list = <fun>

val union : String.t list -> String.t list -> String.t list = <fun>

val free_variables : Syntax.expr -> String.t list -> String.t list = <fun>

val free_variables : Syntax.expr -> String.t list = <fun>

val substitute : Syntax.expr -> string -> Syntax.expr -> Syntax.expr = <fun>
```

<!-- grade_id: cell-d15558191ca8be99, points: 20, grade: true, solution: false -->
```OCaml
assert (alpha_equiv 
          (substitute (parse_string "\\y.x") "x" (parse_string "\\z.z w")) 
          (parse_string "λy.λz.z w"));
assert (alpha_equiv 
          (substitute (parse_string "\\x.x") "x" (parse_string "y"))
          (parse_string "λx.x"));
assert (alpha_equiv 
          (substitute (parse_string "\\x.y") "y" (parse_string "x"))
          (parse_string "λx0.x"))
```

Output:
```
- : unit = ()
```

## Problem 7

Implement a single step of the call-by-value reduction. Implement the function

```ocaml
reduce_cbv : expr -> expr option
```

which does a single step of the call-by-value reduction. Recall that call-by-value reduction is determinisitic. Hence, if reduction is possible, then a single rule applies. `reduce e` returns `Some e'` if reduction is possible and `e'` is the new expression. `reduce e` returns `None` if reduction is not possible.

<!-- grade_id: cell-f0a4289c56a6e0c2, points: 0, grade: false, solution: true -->
```OCaml
let rec mem e l acc = 
    match l with 
    |[] -> acc
    |h::t -> if e=h then true else mem e t acc
   
let mem e l = mem e l false 

let rec rev_list l acc = 
  match l with
  | [] -> acc
  | hd::tl -> rev_list tl (hd::acc)

let rev_list l = rev_list l []

let rec remove e l l1 =
    match l with
    |[] -> rev_list l1
    |h::t -> if e=h then remove e t l1 else (remove e t (h::l1))
    
let remove e l = remove e l []

let rec remove_fun l1 x acc = 
    match l1 with
    |[] -> rev_list acc
    |hd1::tl1 -> if hd1=x then remove_fun tl1 hd1 acc
                 else remove_fun tl1 hd1 (hd1::acc)


let remove_stutter l =
    match l with
    |[]->l
    |hd::t1-> remove_fun t1 hd (hd::[])
    
let rec union l1 l2 = 
    match l2 with 
    |[]->remove_stutter (List.sort String.compare l1)
    |h::t-> union (l1@[h]) t
    
let rec free_variables e l =
    match e with
    |Lam(x,t1) -> remove x (free_variables t1 l)
    |App(t1, t2) -> union (free_variables t1 l) (free_variables t2 l) 
    |Var x -> l@[x]
    
let free_variables e = free_variables e []

let rec substitute expr a b = 
    match expr with 
    |Var x -> if a = x then  b else Var x
    |Lam(x,t1) -> if a=x then Lam(x,t1) else 
         if mem x (free_variables b) = false then Lam(x,substitute t1 a b) 
         else  let w = fresh (x) in substitute (  Lam(w,substitute t1 x (Var w))) a b 
    |App(t1,t2) -> App(substitute t1 a b,substitute t2 a b)
    
let rec is_beta_reducable_cbv e =
    match e with 
    |Var x -> false
    |Lam(x,t1) -> false
    |App(m,n) -> if is_beta_reducable_cbv m = true then true else if  is_beta_reducable_cbv n = true then true 
                else (match m with 
                |Lam(x,m1) -> true 
                |_ -> false )
    
let rec reduce_cbv1 e = 
    match e with
    |App(m,n) -> if is_beta_reducable_cbv m = true then reduce_cbv1 m else if is_beta_reducable_cbv n = true 
                            then reduce_cbv1 n else( match m with 
                                                |Lam(x,t) -> substitute t x n
                                                |_ -> Var "x" )
    |_ -> Var "x"
    
let rec reduce_cbv e =
    match e with
    |App(m,n)-> if is_beta_reducable_cbv m = true then Some (App((reduce_cbv1 m),n)) else 
                if is_beta_reducable_cbv n = true then Some (App(m,(reduce_cbv1 n))) else 
               ( match m with 
                |Lam(x,t) -> Some (substitute t x n)
                |_ -> None )
    |_ -> None
```

Output:
```
val mem : 'a -> 'a list -> bool -> bool = <fun>

val mem : 'a -> 'a list -> bool = <fun>

val rev_list : 'a list -> 'a list -> 'a list = <fun>

val rev_list : 'a list -> 'a list = <fun>

val remove : 'a -> 'a list -> 'a list -> 'a list = <fun>

val remove : 'a -> 'a list -> 'a list = <fun>

val remove_fun : 'a list -> 'a -> 'a list -> 'a list = <fun>

val remove_stutter : 'a list -> 'a list = <fun>

val union : String.t list -> String.t list -> String.t list = <fun>

val free_variables : Syntax.expr -> String.t list -> String.t list = <fun>

val free_variables : Syntax.expr -> String.t list = <fun>

val substitute : Syntax.expr -> string -> Syntax.expr -> Syntax.expr = <fun>

val is_beta_reducable_cbv : Syntax.expr -> bool = <fun>

val reduce_cbv1 : Syntax.expr -> Syntax.expr = <fun>

val reduce_cbv : Syntax.expr -> Syntax.expr option = <fun>
```

<!-- grade_id: cell-653a749cafdc8a78, points: 20, grade: true, solution: false -->
```OCaml
begin match reduce_cbv (parse_string "(\\x.x) ((\\x.x) (\\z.(\\x.x) z))") with
| Some expr -> assert (alpha_equiv expr (parse_string "(λx.x) (λz.(λx.x) z)"))
| None -> assert false
end;
        
begin match reduce_cbv (parse_string "(λx.x) (λz.(λx.x) z)") with
| Some expr -> assert (alpha_equiv expr (parse_string "λz.(λx.x) z"))
| None -> assert false
end;
        
assert (reduce_cbv (parse_string "λz.(λx.x) z") = None);
        
begin match reduce_cbv (parse_string "(λx.y) ((λx.x x) (λx.x x))") with
| Some expr -> assert (alpha_equiv expr (parse_string "(λx.y) ((λx.x x) (λx.x x))"))
| None -> assert false
end;

assert (reduce_cbv (parse_string "x y z") = None)

```

Output:
```
- : unit = ()
```

## Problem 8

Implement a single step of the call-by-name reduction. Implement the function

```ocaml
reduce_cbn : expr -> expr option
```

The rest of the instructions are same as `reduce_cbv`.

<!-- grade_id: cell-97e3a2c41734215b, points: 0, grade: false, solution: true -->
```OCaml
let r = ref 0                                                                       
                                                                                    
let fresh s =                                                                       
  let v = !r in                                                                     
  r := !r + 1;                                                                      
  s ^ (string_of_int v)
  
let rec mem e l acc = 
    match l with 
    |[] -> acc
    |h::t -> if e=h then true else mem e t acc
   
let mem e l = mem e l false 

let rec rev_list l acc = 
  match l with
  | [] -> acc
  | hd::tl -> rev_list tl (hd::acc)

let rev_list l = rev_list l []

let rec remove e l l1 =
    match l with
    |[] -> rev_list l1
    |h::t -> if e=h then remove e t l1 else (remove e t (h::l1))
    
let remove e l = remove e l []

let rec remove_fun l1 x acc = 
    match l1 with
    |[] -> rev_list acc
    |hd1::tl1 -> if hd1=x then remove_fun tl1 hd1 acc
                 else remove_fun tl1 hd1 (hd1::acc)


let remove_stutter l =
    match l with
    |[]->l
    |hd::t1-> remove_fun t1 hd (hd::[])
    
let rec union l1 l2 = 
    match l2 with 
    |[]->remove_stutter (List.sort String.compare l1)
    |h::t-> union (l1@[h]) t
    
let rec free_variables e l =
    match e with
    |Lam(x,t1) -> remove x (free_variables t1 l)
    |App(t1, t2) -> union (free_variables t1 l) (free_variables t2 l) 
    |Var x -> l@[x]
    
let free_variables e = free_variables e []

let rec substitute expr a b = 
    match expr with 
    |Var x -> if a = x then  b else Var x
    |Lam(x,t1) -> if a=x then Lam(x,t1) else 
         if mem x (free_variables b) = false then Lam(x,substitute t1 a b) 
         else  let w = fresh (x) in substitute (  Lam(w,substitute t1 x (Var w))) a b 
    |App(t1,t2) -> App(substitute t1 a b,substitute t2 a b)
    
let rec is_beta_reducable_cbn e =
    match e with 
    |Var x -> false
    |App(Lam(x,m),n) -> true
    |Lam(x,t1) -> false
    |App(m,n) -> is_beta_reducable_cbn m
    
let rec reduce_cbn1 e = 
    match e with 
    |App(Lam(x,m),n) -> substitute m x n
    |App(m,n) -> if is_beta_reducable_cbn m = true then App((reduce_cbn1 m),n) else Var "x"
    |_ -> Var "x"
    
let rec reduce_cbn e =
    match e with 
    |App(Lam(x,m),n) ->Some (substitute m x n)
    |App(m,n) -> if is_beta_reducable_cbn m = true then Some(App((reduce_cbn1 m),n)) else None
    |_ -> None
```

Output:
```
val r : int ref = {contents = 0}

val fresh : string -> string = <fun>

val mem : 'a -> 'a list -> bool -> bool = <fun>

val mem : 'a -> 'a list -> bool = <fun>

val rev_list : 'a list -> 'a list -> 'a list = <fun>

val rev_list : 'a list -> 'a list = <fun>

val remove : 'a -> 'a list -> 'a list -> 'a list = <fun>

val remove : 'a -> 'a list -> 'a list = <fun>

val remove_fun : 'a list -> 'a -> 'a list -> 'a list = <fun>

val remove_stutter : 'a list -> 'a list = <fun>

val union : String.t list -> String.t list -> String.t list = <fun>

val free_variables : Syntax.expr -> String.t list -> String.t list = <fun>

val free_variables : Syntax.expr -> String.t list = <fun>

val substitute : Syntax.expr -> string -> Syntax.expr -> Syntax.expr = <fun>

val is_beta_reducable_cbn : Syntax.expr -> bool = <fun>

val reduce_cbn1 : Syntax.expr -> Syntax.expr = <fun>

val reduce_cbn : Syntax.expr -> Syntax.expr option = <fun>
```

<!-- grade_id: cell-e988e3a110046165, points: 20, grade: true, solution: false -->
```OCaml
begin match reduce_cbn (parse_string "(\\x.x) ((\\x.x) (\\z.(\\x.x) z))") with
| Some expr -> assert (alpha_equiv expr (parse_string "(λx.x) (λz.(λx.x) z)"))
| None -> assert false
end;
        
begin match reduce_cbn (parse_string "(λx.x) (λz.(λx.x) z)") with
| Some expr -> assert (alpha_equiv expr (parse_string "λz.(λx.x) z"))
| None -> assert false
end;
        
assert (reduce_cbn (parse_string "λz.(λx.x) z") = None);
        
begin match reduce_cbn (parse_string "(λx.y) ((λx.x x) (λx.x x))") with
| Some expr -> assert (alpha_equiv expr (parse_string "y"))
| None -> assert false
end;

begin match reduce_cbn (parse_string "(\\x.x x) ((\\z.z) y)") with
| Some expr -> assert (alpha_equiv expr (parse_string "(λz.z) y ((λz.z) y)"))
| None -> assert false
end;

assert (reduce_cbn (parse_string "x y z") = None)

```

Output:
```
- : unit = ()
```

## Problem 9

Implement a single step of the normal order reduction. Implement the function

```ocaml
reduce_normal : expr -> expr option
```

The rest of the instructions are same as `reduce_cbv`.

<!-- grade_id: cell-ba11da6e2b5e2acc, points: 0, grade: false, solution: true -->
```OCaml
let rec mem e l acc = 
    match l with 
    |[] -> acc
    |h::t -> if e=h then true else mem e t acc
   
let mem e l = mem e l false 

let rec rev_list l acc = 
  match l with
  | [] -> acc
  | hd::tl -> rev_list tl (hd::acc)

let rev_list l = rev_list l []

let rec remove e l l1 =
    match l with
    |[] -> rev_list l1
    |h::t -> if e=h then remove e t l1 else (remove e t (h::l1))
    
let remove e l = remove e l []

let rec remove_fun l1 x acc = 
    match l1 with
    |[] -> rev_list acc
    |hd1::tl1 -> if hd1=x then remove_fun tl1 hd1 acc
                 else remove_fun tl1 hd1 (hd1::acc)


let remove_stutter l =
    match l with
    |[]->l
    |hd::t1-> remove_fun t1 hd (hd::[])
    
let rec union l1 l2 = 
    match l2 with 
    |[]->remove_stutter (List.sort String.compare l1)
    |h::t-> union (l1@[h]) t
    
let rec free_variables e l =
    match e with
    |Lam(x,t1) -> remove x (free_variables t1 l)
    |App(t1, t2) -> union (free_variables t1 l) (free_variables t2 l) 
    |Var x -> l@[x]
    
let free_variables e = free_variables e []

let rec substitute expr a b = 
    match expr with 
    |Var x -> if a = x then  b else Var x
    |Lam(x,t1) -> if a=x then Lam(x,t1) else 
         if mem x (free_variables b) = false then Lam(x,substitute t1 a b) 
         else  let w = fresh (x) in substitute (  Lam(w,substitute t1 x (Var w))) a b 
    |App(t1,t2) -> App(substitute t1 a b,substitute t2 a b)
    
let rec is_beta_reducable e =
    match e with 
    |Var x -> false
    |App(Lam(x,m),n) -> true
    |Lam(x,t1) -> is_beta_reducable t1
    |App(m,n) -> if is_beta_reducable m = true then true else is_beta_reducable n
    
let rec reduce_normal1 e =
    match e with 
    |App(Lam(x,m),n) ->  (substitute m x n)
    |App(m,n) -> if is_beta_reducable m = true then App((reduce_normal1 m),n) else 
                    if is_beta_reducable n = true then App(m,(reduce_normal1 n)) else Var "x"
    |Lam(x,m) -> if is_beta_reducable m = true then Lam(x,(reduce_normal1 m)) else Var "x"
    |_ -> Var "x"
    
let rec reduce_normal e = 
    match e with
    |App(Lam(x,m),n) -> Some(substitute m x n)
    |App(m,n) -> if is_beta_reducable m = true then Some(App((reduce_normal1 m),n)) else 
                    if is_beta_reducable n = true then Some(App(m,(reduce_normal1 n))) else None
    |Lam(x,m) -> if is_beta_reducable m = true then Some(Lam(x,(reduce_normal1 m))) else None
    |_ -> None
```

Output:
```
val mem : 'a -> 'a list -> bool -> bool = <fun>

val mem : 'a -> 'a list -> bool = <fun>

val rev_list : 'a list -> 'a list -> 'a list = <fun>

val rev_list : 'a list -> 'a list = <fun>

val remove : 'a -> 'a list -> 'a list -> 'a list = <fun>

val remove : 'a -> 'a list -> 'a list = <fun>

val remove_fun : 'a list -> 'a -> 'a list -> 'a list = <fun>

val remove_stutter : 'a list -> 'a list = <fun>

val union : String.t list -> String.t list -> String.t list = <fun>

val free_variables : Syntax.expr -> String.t list -> String.t list = <fun>

val free_variables : Syntax.expr -> String.t list = <fun>

val substitute : Syntax.expr -> string -> Syntax.expr -> Syntax.expr = <fun>

val is_beta_reducable : Syntax.expr -> bool = <fun>

val reduce_normal1 : Syntax.expr -> Syntax.expr = <fun>

val reduce_normal : Syntax.expr -> Syntax.expr option = <fun>
```

<!-- grade_id: cell-ded05a8d3813e1e9, points: 20, grade: true, solution: false -->
```OCaml
begin match reduce_normal (parse_string "(\\x.x) ((\\x.x) (\\z.(\\x.x) z))") with
| Some expr -> assert (alpha_equiv expr (parse_string "(λx.x) (λz.(λx.x) z)"))
| None -> assert false
end;
        
begin match reduce_normal (parse_string "(λx.x) (λz.(λx.x) z)") with
| Some expr -> assert (alpha_equiv expr (parse_string "λz.(λx.x) z"))
| None -> assert false
end;
        
begin match reduce_normal (parse_string "λz.(λx.x) z") with
| Some expr -> assert (alpha_equiv expr (parse_string "λz. z"))
| None -> assert false
end;
        
begin match reduce_normal (parse_string "(λx.y) ((λx.x x) (λx.x x))") with
| Some expr -> assert (alpha_equiv expr (parse_string "y"))
| None -> assert false
end;

begin match reduce_normal (parse_string "(\\x.x x) ((\\z.z) y)") with
| Some expr -> assert (alpha_equiv expr (parse_string "(λz.z) y ((λz.z) y)"))
| None -> assert false
end;
        
begin match reduce_normal (parse_string "f (\\x.x x) ((\\z.z) y)") with
| Some expr -> assert (alpha_equiv expr (parse_string "f (λx.x x) y"))
| None -> assert false
end;

begin match reduce_normal (parse_string "(\\x.(\\z.z) y) (\\x.x x)") with
| Some expr -> assert (alpha_equiv expr (parse_string "(λz.z) y"))
| None -> assert false
end

```

Output:
```
- : unit = ()
```

```OCaml
let rec eval log depth reduce expr =                     
  if depth = 0 then failwith "non-termination?"                                  
  else                                                                     
    begin match reduce expr with
    | None -> expr
    | Some expr' ->
      if log then print_endline ("= " ^ (string_of_expr expr'));                 
      eval log (depth-1) reduce expr'                                    
    end
  
let eval_cbv = eval true 1000 reduce_cbv
let eval_cbn = eval true 1000 reduce_cbn
let eval_normal = eval true 1000 reduce_normal
```

Output:
```
val eval :
  bool ->
  int -> (Syntax.expr -> Syntax.expr option) -> Syntax.expr -> Syntax.expr =
  <fun>

val eval_cbv : Syntax.expr -> Syntax.expr = <fun>

val eval_cbn : Syntax.expr -> Syntax.expr = <fun>

val eval_normal : Syntax.expr -> Syntax.expr = <fun>
```

<!-- grade_id: cell-8d1c02b4a1ab2c2f, points: 10, grade: true, solution: false -->
```OCaml
let zero = parse_string "\\f.\\x. x" in
let one = parse_string "\\f.\\x. f x" in
let two = parse_string "\\f.\\x. f (f x)" in
let three = parse_string "\\f.\\x. f (f (f x))" in

let plus = parse_string "λm. λn. λs. λz. m s (n s z)" in
let mult = parse_string "λm. λn. λs. λz. m (n s) z" in

assert (alpha_equiv (eval_normal (App (App (plus, one), two))) three);
print_endline "";
assert (alpha_equiv (eval_normal (App (App (mult, one), three))) three);
print_endline "";
assert (alpha_equiv (eval_normal (App (App (mult, zero), three))) zero)
```

Output:
```
= (λn.λs.λz.(λf.λx.f x) s (n s z)) (λf.λx.f (f x))
= λs.λz.(λf.λx.f x) s ((λf.λx.f (f x)) s z)
= λs.λz.(λx.s x) ((λf.λx.f (f x)) s z)
= λs.λz.s ((λf.λx.f (f x)) s z)
= λs.λz.s ((λx.s (s x)) z)
= λs.λz.s (s (s z))

= (λn.λs.λz.(λf.λx.f x) (n s) z) (λf.λx.f (f (f x)))
= λs.λz.(λf.λx.f x) ((λf.λx.f (f (f x))) s) z
= λs.λz.(λx.(λf.λx.f (f (f x))) s x) z
= λs.λz.(λf.λx.f (f (f x))) s z
= λs.λz.(λx.s (s (s x))) z
= λs.λz.s (s (s z))

= (λn.λs.λz.(λf.λx.x) (n s) z) (λf.λx.f (f (f x)))
= λs.λz.(λf.λx.x) ((λf.λx.f (f (f x))) s) z
= λs.λz.(λx.x) z
= λs.λz.z
- : unit = ()
```