Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name below:

```OCaml
let name = "Anumala Venu Madhava Reddy"
let rollno = "CS18B051"
```

Output:
```
val name : string = "Anumala Venu Madhava Reddy"

val rollno : string = "CS18B051"
```

## Important notes about grading:

1. **Compiler errors:** All code you submit must compile. Programs that do not compile will probably receive an automatic zero. If you are having trouble getting your assignment to compile, please visit consulting hours. If you run out of time, it is better to comment out the parts that do not compile, than hand in a more complete file that does not compile.
2. **Late assignments:** Please carefully review the course website's policy on late assignments, as all assignments handed in after the deadline will be considered late. Verify on moodle that you have submitted the correct version, before the deadline. Submitting the incorrect version before the deadline and realizing that you have done so after the deadline will be counted as a late submission.

## Problem 1

Write a function 

```ocaml
cond_dup : 'a list -> ('a -> bool) -> 'a
```

list that takes in a list and preciate and duplicates all elements which satisfy the condition expressed in the predicate. For example, `cond_dup [3;4;5] (fun x -> x mod 2 = 1) = [3;3;4;5;5]`.

<!-- grade_id: cell-7950f1ea0a84b889, points: 0, grade: false, solution: true -->
```OCaml
let rec rev_list l acc = 
  match l with
  | [] -> acc
  | hd::tl -> rev_list tl (hd::acc)

let rev_list l = rev_list l []

let rec cond_dup l f acc  =
     match l with
     |[]-> rev_list acc
     |hd::t1-> if f hd = true then cond_dup t1 f (hd::(hd::acc))
             else cond_dup t1 f (hd::acc)
             
let cond_dup l f = cond_dup l f []
```

Output:
```
val rev_list : 'a list -> 'a list -> 'a list = <fun>

val rev_list : 'a list -> 'a list = <fun>

val cond_dup : 'a list -> ('a -> bool) -> 'a list -> 'a list = <fun>

val cond_dup : 'a list -> ('a -> bool) -> 'a list = <fun>
```

<!-- grade_id: cell-1d772b02ff64382c, points: 10, grade: true, solution: false -->
```OCaml
assert (cond_dup [3;4;5] (fun x -> x mod 2 = 1) = [3;3;4;5;5])
```

Output:
```
- : unit = ()
```

## Problem 2

Write a function 

```ocaml
n_times : ('a -> 'a) * int * 'a -> 'a
```

such that `n_times (f,n,v)` applies `f` to `v` `n` times. For example, `n_times((fun x-> x+1), 50, 0) = 50`. If `n<=0` return `v`.


<!-- grade_id: cell-c28f98bd63f56ba1, points: 0, grade: false, solution: true -->
```OCaml
let rec n_times (f, n, v) = 
   if n <= 0 then v
   else f (n_times (f, n-1, v))
```

Output:
```
val n_times : ('a -> 'a) * int * 'a -> 'a = <fun>
```

<!-- grade_id: cell-98364653c4b77b2d, points: 10, grade: true, solution: false -->
```OCaml
assert (n_times((fun x-> x+1), 50, 0) = 50)
```

Output:
```
- : unit = ()
```

## Problem 3

Write a function 

```ocaml
range : int -> int -> int list 
```

such that `range num1 num2` returns an ordered list of all integers from `num1` to `num2` inclusive. For example, `range 2 5 = [2;3;4;5]`. Raise the exception `IncorrectRange` if `num2 < num1`.

<!-- grade_id: cell-1fd147f7eb3b139d, points: 0, grade: false, solution: true -->
```OCaml
let rec rev_list l acc = 
  match l with
  | [] -> acc
  | hd::tl -> rev_list tl (hd::acc)

let rev_list l = rev_list l []

let rec ran n1 n2 acc =
    if n1>n2 then rev_list acc
    else ran (n1+1) n2 (n1::acc)
    
let rang n1 n2 = ran n1 n2 []

exception IncorrectRange

let range num1 num2 =
    if num1>num2 then raise ( IncorrectRange)
    else rang num1 num2 
```

Output:
```
val rev_list : 'a list -> 'a list -> 'a list = <fun>

val rev_list : 'a list -> 'a list = <fun>

val ran : int -> int -> int list -> int list = <fun>

val rang : int -> int -> int list = <fun>

exception IncorrectRange

val range : int -> int -> int list = <fun>
```

<!-- grade_id: cell-bba984120a807d42, points: 10, grade: true, solution: false -->
```OCaml
assert (range 2 5 = [2;3;4;5])
```

Output:
```
- : unit = ()
```

## Problem 4

Write a function 

```ocaml
zipwith : ('a -> 'b -> 'c) -> 'a list -> 'b list -> 'c list
```

such that `zipwith f l1 l2` generates a list whose `ith` element is obtained by applying `f` to the `ith` element of `l1` and the `ith` element of `l2` . If the lists have different lengths, the extra elements in the longer list should be ignored. For example, `zipwith (+) [1;2;3] [4;5] = [5;7]`.

<!-- grade_id: cell-932054b41dace011, points: 0, grade: false, solution: true -->
```OCaml
let rec rev_list l acc = 
  match l with
  | [] -> acc
  | hd::tl -> rev_list tl (hd::acc)

let rev_list l = rev_list l []

let rec zipwith f l1 l2 l =
    match (l1, l2) with
    |([],_)->rev_list l
    |(_,[])->rev_list l
    |(hd1::t1,hd2::t2)-> zipwith f t1 t2 ((f hd1 hd2)::l)
    
let zipwith f l1 l2 = zipwith f l1 l2 []
```

Output:
```
val rev_list : 'a list -> 'a list -> 'a list = <fun>

val rev_list : 'a list -> 'a list = <fun>

val zipwith : ('a -> 'b -> 'c) -> 'a list -> 'b list -> 'c list -> 'c list =
  <fun>

val zipwith : ('a -> 'b -> 'c) -> 'a list -> 'b list -> 'c list = <fun>
```

<!-- grade_id: cell-55aefe5fa4f927d0, points: 10, grade: true, solution: false -->
```OCaml
assert (zipwith (+) [1;2;3] [4;5] = [5;7]) 
```

Output:
```
- : unit = ()
```

## Problem 5

Write a function 

```ocaml
buckets : ('a -> 'a -> bool) -> 'a list -> 'a list list 
```

that partitions a list into equivalence classes. That is, buckets equiv lst should return a list of lists where each sublist in the result contains equivalent elements, where two elements are considered equivalent if equiv returns true. For example:


```ocaml
buckets (=) [1;2;3;4] = [[1];[2];[3];[4]]
buckets (=) [1;2;3;4;2;3;4;3;4] = [[1];[2;2];[3;3;3];[4;4;4]]
buckets (fun x y -> (=) (x mod 3) (y mod 3)) [1;2;3;4;5;6] = [[1;4];[2;5];[3;6]]
```

The order of the buckets must reflect the order in which the elements appear in the original list. For example, the output of `buckets (=) [1;2;3;4]` should be `[[1];[2];[3];[4]]` and not `[[2];[1];[3];[4]]` or any other permutation.

The order of the elements in each bucket must reflect the order in which the elements appear in the original list. For example, the output of `buckets (fun x y -> (=) (x mod 3) (y mod 3)) [1;2;3;4;5;6]` should be `[[1;4];[2;5];[3;6]]` and not `[[4;1];[5;2];[3;6]]` or any other permutations.

Assume that the comparison function `('a -> 'a -> bool)` is commutative, associative and idempotent.

Just use lists. Do not use sets or hash tables.

List append function `@` may come in handy. `[1;2;3] @ [4;5;6] = [1;2;3;4;5;6]`.

<!-- grade_id: cell-f4d88f256cc9db8b, points: 0, grade: false, solution: true -->
```OCaml
let find_class l1 f1 y =
    match l1 with 
    |[]->false
    |h::t-> f1 y h
    
let rec add_to_class l2 f x =
    match l2 with
    |[]->l2@[[x]]
    |hd::tl-> if (find_class hd f x) = true then [hd@[x]]@tl else [hd]@( add_to_class tl f  x)
    
let rec buckets p l l4 =
    match l with 
    |[]->l4
    |h::t-> buckets p t (add_to_class l4 p h)
    
    
let buckets p l = buckets p l []
```

Output:
```
val find_class : 'a list -> ('b -> 'a -> bool) -> 'b -> bool = <fun>

val add_to_class : 'a list list -> ('a -> 'a -> bool) -> 'a -> 'a list list =
  <fun>

val buckets : ('a -> 'a -> bool) -> 'a list -> 'a list list -> 'a list list =
  <fun>

val buckets : ('a -> 'a -> bool) -> 'a list -> 'a list list = <fun>
```

<!-- grade_id: cell-fdbd3b5564e151cb, points: 20, grade: true, solution: false -->
```OCaml
assert (buckets (=) [1;2;3;4] = [[1];[2];[3];[4]])
; assert (buckets (=) [1;2;3;4;2;3;4;3;4] = [[1];[2;2];[3;3;3];[4;4;4]])
; assert (buckets (fun x y -> (=) (x mod 3) (y mod 3)) [1;2;3;4;5;6] = [[1;4];[2;5];[3;6]])
```

Output:
```
- : unit = ()
```

## Problem 6

Write a function 

```ocaml
remove_stutter : 'a list -> 'a list
```

that removes stuttering from the original list. For example, `remove_stutter [1;2;2;3;1;1;1;4;4;2;2]= [1;2;3;1;4;2]`. Option type may come in handy. 

<!-- grade_id: cell-8865e0dd05cbd32d, points: 0, grade: false, solution: true -->
```OCaml
let rec rev_list l acc = 
  match l with
  | [] -> acc
  | hd::tl -> rev_list tl (hd::acc)

let rev_list l = rev_list l []

let rec remove_fun l1 x acc = 
    match l1 with
    |[] -> rev_list acc
    |hd1::tl1 -> if hd1=x then remove_fun tl1 hd1 acc
                 else remove_fun tl1 hd1 (hd1::acc)


let remove_stutter l =
    match l with
    |[]->l
    |hd::t1-> remove_fun t1 hd (hd::[])
```

Output:
```
val rev_list : 'a list -> 'a list -> 'a list = <fun>

val rev_list : 'a list -> 'a list = <fun>

val remove_fun : 'a list -> 'a -> 'a list -> 'a list = <fun>

val remove_stutter : 'a list -> 'a list = <fun>
```

<!-- grade_id: cell-aa3fb845581c6fce, points: 15, grade: true, solution: false -->
```OCaml
assert (remove_stutter [1;2;2;3;1;1;1;4;4;2;2] = [1; 2; 3; 1; 4; 2]) 
```

Output:
```
- : unit = ()
```

## Problem 7

Write a function

```ocaml
flatten : 'a list list -> 'a list
```

that flattens a list. For example, `flatten [[1;2];[3;4]] = [1;2;3;4]`.

<!-- grade_id: cell-1ba30bbe74fd3005, points: 0, grade: false, solution: true -->
```OCaml
let rec rev_list l acc = 
  match l with
  | [] -> acc
  | hd::tl -> rev_list tl (hd::acc)

let rev_list l = rev_list l []

let rec add_list l1 acc=
    match l1 with 
    |[]-> acc
    |h1::t1 -> add_list t1 (h1::acc)

let rec flatten l acc=
    match l with
    |h::t -> flatten t (add_list h acc)
    |[]->rev_list acc
    
    
let flatten l = flatten l []

```

Output:
```
val rev_list : 'a list -> 'a list -> 'a list = <fun>

val rev_list : 'a list -> 'a list = <fun>

val add_list : 'a list -> 'a list -> 'a list = <fun>

val flatten : 'a list list -> 'a list -> 'a list = <fun>

val flatten : 'a list list -> 'a list = <fun>
```

<!-- grade_id: cell-8e986c3f5dbf5e89, points: 10, grade: true, solution: false -->
```OCaml
assert (flatten ([[1;2];[3;4]]) = [1;2;3;4])
```

Output:
```
- : unit = ()
```

## Problem 8

Write a function

```ocaml

type 'a tree = Leaf | Node of 'a tree * 'a * 'a tree

fold_inorder : ('a -> 'b -> 'a) -> 'a -> 'b tree -> 'a
```

That does a inorder fold of the tree. For example,

`fold_inorder (fun acc x -> acc @ [x]) [] (Node (Node (Leaf,1,Leaf), 2, Node (Leaf,3,Leaf))) = [1;2;3]`

<!-- grade_id: cell-7170388040d2ef4b, points: 0, grade: false, solution: true -->
```OCaml
type 'a tree = Leaf | Node of 'a tree * 'a * 'a tree

let rec rev_list l acc = 
  match l with
  | [] -> acc
  | hd::tl -> rev_list tl (hd::acc)

let rev_list l = rev_list l []

let rec modify t l =
    match t with 
    |Leaf ->  l
    |Node(t1,x,t2) -> modify t2 (x::(modify t1 l))

let help t = rev_list(modify t [])

let rec fold_inorder f acc l = 
    match l with 
    |[]-> acc
    |h::t-> fold_inorder f (f acc h) t

let fold_inorder f acc t = fold_inorder f acc (help t)
```

Output:
```
type 'a tree = Leaf | Node of 'a tree * 'a * 'a tree

val rev_list : 'a list -> 'a list -> 'a list = <fun>

val rev_list : 'a list -> 'a list = <fun>

val modify : 'a tree -> 'a list -> 'a list = <fun>

val help : 'a tree -> 'a list = <fun>

val fold_inorder : ('a -> 'b -> 'a) -> 'a -> 'b list -> 'a = <fun>

val fold_inorder : ('a -> 'b -> 'a) -> 'a -> 'b tree -> 'a = <fun>
```

<!-- grade_id: cell-8bd22519c4af64bb, points: 15, grade: true, solution: false -->
```OCaml
assert (fold_inorder (fun acc x -> acc @ [x]) [] (Node (Node (Leaf,1,Leaf), 2, Node (Leaf,3,Leaf))) = [1;2;3]) 
```

Output:
```
- : unit = ()
```

## Problem 9

The usual recursive formulation of fibonacci function 

```ocaml
let rec fib n = 
  if n = 0 then 0
  else if n = 1 then 1
  else fib (n-1) + fib (n-2)
```

has exponential running time. It will take a long time to compute `fib 50`. You might have to interrupt the kernel if you did try to do `fib 50` in the notebook.

But we know that fibonacci number can be computed in linear time by remembering just the current `cur` and the previous `prev` fibonacci number. In this case, the next fibonacci number is computed as the sum of the current and the previous numbers. Then the program continues by setting `prev` to be `cur` and `cur` to be `cur + prev`.

Implement a tail recursive function `fib_tailrec` that uses this idea and computes the `nth` fibonacci number in linear time. 

```ocaml
fib_tailrec : int -> int
```

<!-- grade_id: cell-39286aaf6b703959, points: 0, grade: false, solution: true -->
```OCaml
let rec fib_tailrec n x curr prev =
    if x=n then curr
    else if x=(n+1) then prev
    else fib_tailrec n (x+1) (curr+prev) curr 


let fib_tailrec n = fib_tailrec n 1 1 0 
```

Output:
```
val fib_tailrec : int -> int -> int -> int -> int = <fun>

val fib_tailrec : int -> int = <fun>
```

<!-- grade_id: cell-57d3e7608b2286f5, points: 15, grade: true, solution: false -->
```OCaml
assert (fib_tailrec 50 = 12586269025)
```

Output:
```
- : unit = ()
```